# Lab 04

In Tuesday's class you had learnt about how to infer parameters of, and draw conclusions about a population using a statistic (number/property describing a characteristic of a sample). 
Further, there were examples of wine-data which provided information about:
* Understanding sampling distributions
* Method 1: Point estimates
* Method 2: Confidence intervals
* Method 3: Hypothesis testing

In this lab session, we will look into the following
* Exercise importing function for files of .txt, .xlsx, .csv 
* Demostration of Central Limit Theorem.
* Confidence Intervals. 
* t-distributions and stats packages. 

In some point, we will make use of random numbers as data and/or wine-data. 

## Lab Setup

```python
import pandas as pd
import numpy as np
rng_seed = 42
import seaborn as sns
from scipy import stats 
from scipy.stats import norm
import zipfile
import sys
import matplotlib.pyplot as plt
import matplotlib as mpl
%matplotlib inline

import os
if os.environ["HOME"]=='/home/jovyan':
    !pip install --upgrade otter-grader
    
import otter
grader = otter.Notebook()
```

```python
# set up random number generator 
rng = np.random.default_rng(rng_seed)
```

## Examples 

### Importing datafiles



#### Importing Excel files
For importing the files as with multiple sheet in excel workbook, in Python 3.0+, you can use `pd.read_excel()`.

For example, if workbook.xlsx has sheets 'Sheet1' and 'Sheet2', then use `pd.read_excel('workbook.xlsx', 'Sheet1')` or  `('workbook.xlsx', 'Sheet2')`. 



#### Importing from Zipped files 

Now, try importing csv file present inside the zip file named 'avocado' directly (make use of 'zipfile' module).
* https://docs.python.org/3/library/zipfile.html

```python
# Example code of how to load in data stored within a zip-file
# import zipfile  # already imported above 
zf = zipfile.ZipFile('data/avocado.zip')
avacado = pd.read_csv(zf.open('avocado.csv'), index_col = 0)
avacado.head()
```

### System information 

If you ever need to check which version of Python or what version of a module you are working with you can use something similar to the following code. 

```python
# Get to know your system versions
# import sys  # already imported above
print('Python version:', sys.version, '\n')
print('Pandas version:', np.__version__ , '\n')
print('Seaborn version: ', sns.__version__, '\n')
```

Let's load the wine data again.

```python
# Load wine data 
url  =  "https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/winequality-red.csv" 
wine = pd.read_csv(url, sep=";")
wine.head()
```

```python
wine.shape
```

# Inferential Statistics

The piece of code below is adapted from `nb.week4.ipynb` class material. A sample of size 5 is selected and the mean is calculated.  This is repeated 1000 times, collecting the sample means. 
A histogram function is invoked for sampled means. 

Observe the distribution of the mean samples - is it normal, uniform, gamma, beta ..... ?

```python
'''The piece of code from nb.week3.ipynb is provided, use it to  
compute the mean for the 'n' random samples and store in array, means[i]
This is repeated 'N_tests' times. '''

mu = wine["pH"].mean()                                # Mean of population
print("mean pH, mu = " + str(mu))

N_tests = 1000                                        # Num of test to perform
n = 5                                                 # number of samples
means = [0] * N_tests

for i in range(N_tests):
    observations = rng.choice(wine.index.values,n)
    sampled_wines = wine.loc[observations]
    means[i] = sampled_wines["pH"].mean()
    
# Histogram in Matplotlib    
nbins = 15
_ = plt.hist(means, nbins, color= 'red' )
plt.title(' Histogram for sampled mean of pH')
plt.xlabel('pH')
plt.ylabel('Counts')

# Histogram in seaborn
sns.set_style('whitegrid')
sns.set_style('ticks')
fig, ax = plt.subplots()
fig.set_size_inches(8, 6)
sns.histplot(means, bins=nbins, color ='red', kde=True)
plt.title('Histogram for sampled mean of wine[ph] distribution')
plt.xlabel('pH')

plt.show()
```

### Central Limit Theorem

From Chap 4, we had a chance to look into Central Limit Theorem. The initial population distribution was not observed in the class and the mean sampled distribution, i.e., mean of `wine['pH']`, evidently displayed normal distribution (see the plots above).



<!-- BEGIN QUESTION -->

## Exercise 1

Check what was the distribution of the initial population ? Create the  Histogram and read the discussion below.

https://seaborn.pydata.org/generated/seaborn.histplot.html#seaborn.histplot

```python
# Plot the distribution of the population itself i.e., wine['pH']. 
#  Use seaborn rather than matplotlib to create the histogram

...

```

<!-- END QUESTION -->

Since, the population distribution was also normal distribution plot. We might expect that the sampled mean could be normal distribution. What if we try with different population distribution using random generator available in python library and observe the profile. Does it agree with Central Limit Theorem?

For options regarding random numbers generation in certain distribution pattern, see the link:  
https://docs.scipy.org/doc/numpy-1.14.0/reference/routines.random.html

*Think about:* What if the initial population is other than normal distribution. Would CLT come into picture? 



## Exercise 2 - Uniform distribution

In Exercise 1, the distribution of wine pH values is approximately normal.  What about if this data follows another distribution?  How does it affect the distribution of the sampled means?

Let's look if the initial data came from a Uniform distribution.

### Exercise 2A - Create distribution

Create a DataSeries sampling 1000 numbers from the uniform distribution. 

```python
# Create a pandas DataSeries - created using random numbers 
#  Sample 1000 numbers from the uniform distribution - unif(0,1) 
#  Hint: look at using the random number generator
#   seeded above:  rng.uniform()

s = pd.Series(rng.uniform(...))

s.head()
```

```python
grader.check("q2a")
```

### Exercise 2B - Histogram 

Create a histogram of the samples stored in `s` above (the values sampled from the uniform distribution). 

<!-- BEGIN QUESTION -->



```python
# Plot a histogram for the initial population, stored in 's'
#  Use seaborn for the plot 

...

```

<!-- END QUESTION -->

### Exercise 2C - Histogram of Sampled Means

Create a histogram of the sampled means. 

<!-- BEGIN QUESTION -->



```python
# In this cell, collection of 'n' random samples 'N-test' times 
#  storing the mean of each collection in the array.  
# Then, create a histogram of the sampled means. 

# As a start I have copied the most of the code for where this was done for wine-pH.  
# Modify it to get the sample means of the uniformly distributed sample population.

mu = ...
print("mean population, mu = " + str(mu))

N_tests = 1000
n = 50 
means = [0] * N_tests

for i in range(N_tests):
    observations = rng.choice(s.index.values,n)
    sampled_pop = s.loc[observations]
    means[i] = sampled_pop.mean()
    
    
nbins = 15
# Create a histogram of the sampled means from the uniform distribution

...

```

<!-- END QUESTION -->

#####  Once the code is up and running. Examine the following for better understanding. 
 * How is plot sensitive for 'n' - no of samples drawn from population?
 * How does the profile change when 'N_test' - is increased or decreased?

## Exercise 3 - Gamma distributions

Let's examine this again if the initial data is from a Gamma distributions. 

Will the Central Limit Theorem agree here as well? 

Read through Chapter 4 for more information.  

```python

# To try Gamma distribution function - np.random.gamma
#   with the random number generator - rng.gamma

s = pd.Series( )

 
```

```python
grader.check("q3a")
```

<!-- BEGIN QUESTION -->

### Exercise 3B - Histogram 

Create a histogram of the samples stored in `s` above (the values sampled from the gamma distribution). 

```python
# Histogram plot of initial population
#  Use seaborn for the plot

...

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Exercise 3C - Histogram of Sampled Means

Create a histogram of the sampled means.

```python
# Perform the similar operation as above but on a different distribution, e.g., gamma. 
# Refer to the link  provided earlier to generate the initial samples. 
# Continue to use seaborn for plotting.

mu = s.mean();
print("mean of random dataset, mu = " + str(mu))

N_tests = 1000
n = 50
means = [0] * N_tests





# Histogram of sampled means 
 
 
...

```

<!-- END QUESTION -->

## Exercise 4 - Wald distributions

Let's examine this again if the initial data is from a Wald distributions. 

Will the Central Limit Theorem agree here as well? 

Read through Chapter 4 for more information.  

```python

# Perform the similar operation as above but on a different distribution, e.g., wald. 
# Refer to the link  provided earlier to generate the initial samples. 
# Continue to use seaborn for plotting.


# To try Wald distribution function - np.random.wald
#  use the random number generator - rng.wald 

s = pd.Series( )

 
```

```python
grader.check("q4a")
```

<!-- BEGIN QUESTION -->

### Exercise 4B 

Repeat the analysis from above using the Wald distribution. 

```python
# Histogram plot of initial population
#  Use seaborn for the plot

...

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Exercise 4C 

Repeat the analysis from above. 

```python
# Perform the similar operation as above but on a different distribution, e.g., wald. 
# Refer to the link  provided earlier to generate the initial samples. 
# Continue to use seaborn for plotting.

mu = s.mean();
print("mean of random dataset, mu = " + str(mu))

N_tests = 1000
n = 50
means = [0] * N_tests





# Histogram of sampled means 
 
...

```

<!-- END QUESTION -->

## Measuring Variability of Estimates

Since, both the uniform and other distribution in population resulted in the normal (Gaussian) distribution type in the sampled mean. Now, Method-1 (point estimates) and Method 2 (Confidence Interval) can be applied.  



<!-- BEGIN QUESTION -->

## Exercise 5 - Confidence Intervals

Refer to section 4.3.1 and 4.3.2 in the textbook and from `nb.week4` notebook. 

**Q** Explain in words (in the below Markdown cell), what is a Confidence Interval? 

**Q** What does it mean 95% C.I ? 

**ANSWER HERE:**

* A confidence interval is  ...

* 95% CI ... 

<!-- END QUESTION -->

### Example - Confidence Intervals

Spend some time reviewing https://docs.scipy.org/doc/scipy/reference/stats.html


* Below piece of code is replicated from `nb.week4`. Have a look into the calculation of Confidence Intervals. What happens when confidence level is changed? to 90% --> 95% --> 99% --> 99.9 %.

* Check with different no of samples. 
* Check what happens if the no. of N-test is increased. 

Read: 4.3.2 in textbook

```python
''' observe the output plot while changing the Confidence level, 
no of samples and no of times the test is performed '''

# some sampled mean. Note, not mean of population
xbar = sampled_wines["pH"].mean() 
# some sampled S.D. Note, not s.d. of population
sigma = sampled_wines["pH"].std() 


alpha = 0.01

z_alphadiv2 = norm.ppf(1-alpha/2)
confidence = norm.cdf(z_alphadiv2)-norm.cdf(-z_alphadiv2)
ci = [xbar - sigma*z_alphadiv2, xbar + sigma*z_alphadiv2 ]
print ("We are", round (confidence,5)*100 , "% confident that the interval,")
print ("  ",  ci,  "contains the mean.")

N_test = 100
n = 25

means = np.array([0.0] * N_test)   # initialize list for means
sigma = np.array([0.0] * N_test)   # initialize list for standard deviation
ci = np.array([[0.0,0.0]] * N_test)  # initialize list for confidence intervals
mu = wine["pH"].mean()  # true mean


for i in range(N_test):
    observations = rng.choice(wine.index.values,n)
    sampled_wines = wine.loc[observations]
    means[i] = sampled_wines["pH"].mean()
    sigma[i] = sampled_wines["pH"].std()
    ci[i] = means[i] + np.array([-sigma[i] * z_alphadiv2/np.sqrt(n), 
                                 sigma[i]*z_alphadiv2/np.sqrt(n)])

out1 = ci[:,0] > mu # flag CI that do not contain the "true" mean
out2 = ci[:,1] < mu # flag CI that do not contain the "true" mean


fig, ax = plt.subplots(1, 1, figsize=(8, 4))
ind = np.arange(1, N_test+1)
ax.axhline(y = mu, 
           xmin = 0, 
           xmax = N_test+1, 
           color = [0, 0, 0])

ci = np.transpose(ci)
ax.plot([ind,ind], 
        ci, 
        color = '0.75', 
        marker = '_', 
        ms = 0, 
        linewidth = 3)
ax.plot([ind[out1],ind[out1]], 
        ci[:, out1], 
        color = [1, 0, 0, 0.8], 
        marker = '_', 
        ms = 0, 
        linewidth = 3)
ax.plot([ind[out2],ind[out2]], 
        ci[:, out2], 
        color = [1, 0, 0, 0.8], 
        marker = '_',
        ms = 0, 
        linewidth = 3)
ax.plot(ind, 
        means, 
        color = [0, .8, .2, .8], 
        marker = '.',
        ms = 10, 
        linestyle = '')
ax.set_ylabel("Confidence interval for the samples' mean estimate",
              fontsize = 12)
ax.set_xlabel('Samples (with %d observations). '  %n, 
              fontsize = 12)
n = plt.show()
```

## Exercise 6

Now assume that only the good quality wines are distilled at one of the wineries. It is known they contain pH greater than and equal to 3.5. The sampled mean by a quality inspector was 3.56 for a batch of 15 such wines. What is the 95% CI for mean? 
 

```python


# We are interested in only a wines which have pH greater than 3.5. 
# This becomes our population  of interest
rng = np.random.default_rng(rng_seed)

n = 
alpha = 

# Sample Mean
Xbar = 3.56

# Filtering data for higher quality wines with pH >= 3.5 
highquality = 
sigma = highquality['pH'].std()


observations = np.random.choice(highquality.index.values, 15)
sampled_observations = highquality.loc[observations]

mu = sampled_observations['pH'].mean()
sigma = sampled_observations['pH'].std()
print('The mean is :', mu, '& the Std deviation is:', sigma)


# standard error computation
se = sigma/np.sqrt(n)
z_alphadiv2 = norm.ppf(1-alpha/2)
confidence = norm.cdf(z_alphadiv2)-norm.cdf(-z_alphadiv2)
CI = [Xbar - se*z_alphadiv2,Xbar + se*z_alphadiv2 ]

print('We are', round(100*confidence,3), "% confident that the interval")
print("   ", CI , "contains the mean.")

```

```python
grader.check("q6")
```

## Exercise 7

Using `stats.norm.interval()`, CI can be computed.  Use this function to compute the 95% CI for mean of 3.56 (batch of 15)

```python


CIs = stats.norm.interval(...)
print('We are', round(100*confidence,3), "% confident that the interval")
print('   ', CIs , "contains the mean.")
       
```

```python
grader.check("q7")
```

### Example - t distribution


If number of samples are low, it is better to use the t distribution.

Recall, the exercise from `nb.week4`.  An engineer working for Ford is interested in the population of all vehicles that have an engine size of 3.0L or larger, and is particular interested in $\mu$, the highway mileage (mpg). Assume the population is normally distributed. The sample mean among a random sample of 14 vehicles is 18.3 mpg, and the sample standard deviation is 5.1 mpg (note: $\sigma$ is unknown). What is the 95% CI for $\mu$?

```python
from scipy.stats import t
alpha = 0.05
xbar = 18.3
n = 14
dof = n-1
se = 5.1/np.sqrt(n)
t_alphadiv2 = t.ppf(1-alpha/2,dof)
confidence = t.cdf (t_alphadiv2,dof)-t.cdf(-t_alphadiv2, dof)
ci = [xbar - se*t_alphadiv2, xbar + se*t_alphadiv2]


print('we are', round(100*(1-alpha), 3), '%', "confident that the interval")
print('   ',  ci , "contains the mean.")


```

## Exercise 8 

We can make a function out of the above code and/or can invoke Stats package from SciPy for the same. Recommended is to get use to stats package - `scipy.stats.t()`

Apply `stats.t.interval()` function here to obatin the 95% CI, the sigma for this is known to be around 5.1. `xbar` is 18.3. The degrees of freedom are same as above n-1. 

https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.t.html


```python
# Use scipy.stats  
from scipy.stats import t
CIt = stats.t.interval(...)
print('we are', round(100*(1-alpha), 3), '%', "confident that the interval")
print("   ", CIt , "contains the mean.")
```

```python
grader.check("q8")
```

### Example 

Calculate the standard error using `scipy.stats.sem()`

https://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.stats.sem.html

Standard error is computed by \begin{align}
S.E = \frac{\sigma}{\sqrt{n}}
\end{align}



```python
from scipy import stats
alpha = 0.05
sigma = sampled_observations['pH'].std()
n = 15
se = sigma/np.sqrt(n)
sef =  stats.sem(sampled_observations['pH'])
print('S.E from sigma and n is: ', se)
print(' and S.E from stats is also: ', sef)
```

## Exercise 9 

Now you have worked on `stats.normal.interval` and `stats.t.interval` to compute the CI. Try to solve the questions below.

 The random daily quality check at packaging unit yielded sigma of 2.0 for a samples of 150 observations and the mean was 3.7. Find the 90% confidence interval for of the mean.

```python
# compute the S.E (standard Error) using the formula explained earlier. 

# Use the most appropriate function: norm.interval or t.interval 
alpha = ...
# CIs = stats.norm.interval(...)
# CIs = stats.t.interval(...)

print('we are', round(100*(1-alpha), 3), "% confident that the interval")
print("   ", CIs , "contains the mean.")
```

```python
grader.check("q9")
```

### Extra Exercises 

Now, if you have time, try to implement what you have learnt in this session to the avacado dataset, we had extracted from zipfile.

1. What is the distribution of average price in the dataset?
2. Distribution of sampled price?
3. CI for a case. 

## 🎉 Congratulations! You Have Completed Lab 04

### 📤 Submission Instructions

Below, you will find a code cell that will automatically generate a **ZIP file** containing your autograded answers and submit it to the **Lab 04** assignment on **Gradescope**.

Before running the submission cell, please ensure that:
- ✅ **All notebook cells have been executed in order** so that all outputs (including figures and graphs) are visible.
- 💾 **Your notebook is saved** before exporting and submitting.

If you experience any issues during submission, please contact the course staff **well in advance of the deadline** for assistance.

Once everything is confirmed, proceed to run the submission cell below.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2a (points: 2)
```
>>> np.isclose(s.iloc[0], 0.033877795, rtol=1e-06, atol=1e-09)
np.True_
```
```
>>> np.isclose(s.iloc[5], 0.053325084, rtol=1e-06, atol=1e-09)
np.True_
```


Tests for q3a (points: 2)
```
>>> np.isclose(s[0], 3.627469487)
np.True_
```
```
>>> np.isclose(s[10], 4.45165017)
np.True_
```


Tests for q4a (points: 2)
```
>>> np.isclose(s[0], 2.883868)
np.True_
```
```
>>> np.isclose(s[10], 1.09204396)
np.True_
```


Tests for q6 (points: [1, 1, 1, 1])
```
>>> np.isclose(mu, 3.5633333)
np.True_
```
```
>>> np.isclose(sigma, 0.06229729)
np.True_
```
```
>>> np.isclose(round(100 * confidence, 3), 95.0)
np.True_
```
```
>>> type(CI[0]) == np.float64
True
```


Tests for q7 (points: 4)
```
>>> type(CIs[0]) == np.float64
True
```
```
>>> type(CIs) == tuple
True
```


Tests for q8 (points: [1])
```
>>> type(CIt[0]) == np.float64
True
```


Tests for q9 (points: [1])
```
>>> type(CIs[0]) == np.float64
True
```
