# Lab 1

## UN 5550, Spring 2026

## 👋 Welcome to the First Lab: Python Basics

Welcome to the first lab section with Python!

Labs in this course are designed to **complement the homework assignments**. While labs are submitted and graded, you are strongly encouraged to work through **all questions carefully**. The skills and concepts you practice here will be essential for your homework assignments and future coursework. Remember, you learn data science by **_doing_** data science.

### 🎯 Goals of This Lab
By the end of this lab, you will be able to:
- Access and use the core technologies required for this course
- Navigate and work effectively in the Jupyter Notebook environment
- Apply basic arithmetic operators in Python
- Define and work with Python names (variables)


```python
# Initialize Otter (the autograding system)
import numpy as np 
import os

if os.environ["HOME"]=='/home/jovyan':
    !pip install --upgrade otter-grader

import otter
grader = otter.Notebook()
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Jupyter Notebook

# Cells

**Not these!** &rarr;

<div align=center>
    
<img src='images/cell_graphic.jpeg' width=200>
   
</div>

<br>

Jupyter Notebooks are made up of cells. There are two types of cells in a Jupyter Notebook: code cells and [Markdown](https://www.markdownguide.org/getting-started/) cells. 
- Code cells are where we write our Python code.
- Markdown cells allow us to write text, like the text you're reading right now. In Homework 1, you'll get a chance to learn a bit of Markdown.

## Code cells

Running a code cell will execute all of the code it contains.  Any output produced by the code will appear below the cell. Notice the `In:` and `Out:` statements to the left of the cells, followed by a number. That number designates the order in which the cell was run. For example, the number `[3]` next to a cell would indicate that this cell is the third cell that has been run in the notebook so far, or that this is the third time you've run this cell if you've re-run the cell multiple times.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle. Next, do one of the following:
* **Keyboard shortcut**: Hold down the `Shift` key and press `Return` or `Enter`.
* **Button click**: Click the "Run" button in the cell toolbar above.

While the above two options give no difference in computer runtime performance, it's useful to internalize some common keyboard shortcuts early in your Jupyter experience (think of it like learning "Ctrl-C" or "Command-C" to copy instead of "Right click -> Copy to Clipboard"). While they're not required for this course, keyboard shortuts will greatly improve your quality of life when programming Python (just trust us on this).


We can use code cells to do arithmetic in Python. Try running the code cells below:

```python
3 + 4 ** 2
```

The order of evaluation is what you might expect: parentheses, then exponentiation, then multiplication and division, then addition and subtraction; from left to right. 

```python
2 *(3 + 10)
```

You can use parentheses to modify the typical order of evaluation. **How would the above output change, if you removed the parentheses?**

```python
4 ** 2
1 + 2 + 4 + 8 + 16 + 32 + 64
```

**Note that only the value from evaluating the last expression is displayed;** the output from `4 **2` isn't shown at all. This is expected behavior.

```python
10000
```

Also notice that this cell contains a very simple expression that evaluates to `1000`.

```python
10 + 20 + 30 - max(10, 20, 30) - min(10, 20, 30)
```

The output you see from running these five code cells are the values `19`, `16`, `127`, `1000` and `20` displayed, respectively. 

Comments in code cells are denoted using a hashtag (`#`). Comments are used to explain what code does when necessary (or, occasionally in our case, to provide instructions).

```python
# This is a comment!
19 - 18
```

If you try and use syntax Python doesn't know about, it will give you an error.

```python
3 + (2 * 4)
```

```python
(3 + 2) * 4
```

<!-- BEGIN QUESTION -->

<br/>

---

## Question 1: Markdown Cells


To edit an existing Markdown cell, double click it and change whatever you want. To *render* (display) it, you can "run" it like you'd run a code cell.

**Edit the following cell to contain your favorite color.**

*Note:* This is how questions will be formatted in your homework assignments.

My favorite color is *(Type your answer here, replacing the text)*

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 2: Converting cells

To convert a cell from code to Markdown (or vice versa), you should use the dropdown menu in the cell toolbar. For instance, to convert the following cell from Markdown to code, you should:
1. Click the cell.
2. Click "Markdown" in the cell toolbar, and change it to "Code."

Your cell toolbar should then look like this:

<img src="images/code-celltoolbar.png" width="50%"/>

**Change the following cell from a Markdown cell to a code cell**, and run it by pressing the "Run" button or using the keyboard shortcut `Shift + Enter`.

3 + 4 + 5  # Change this to a code cell

<!-- END QUESTION -->

---

## Adding and deleting cells

To add cells, you can go to "Insert" in the toolbar at the top, then click "Insert cell Above" or "Insert cell Below." To delete cells, you can go to "Edit" in the toolbar at the top and click "Delete cells". (Be very careful when deleting a cell.)

For your convenience, we provide this screenshot that details how to switch between "Edit mode" and "Command mode", which is relevant if you want to learn keyboard shortcuts. Again, keyboard shortcuts are not required for this course, but learning some now will greatly improve your programming quality of life.

<img src='images/edit-mode.png' width=400>

Some useful keyboard shortcuts:

| Action | Mode | Keyboard shortcut |
| :--- | :--- | :---
| Run cell + jump to next cell | Either (puts you in edit mode) | `SHIFT + ENTER` |
| Save notebook | Either | `CTRL/CMD + S` |
| Switch to command mode | Either (puts you in command mode) | `ESCAPE` |
| Switch to edit mode | Command | `ENTER` | 
| Comment out the current line | Edit | `CTRL/CMD + /` |
| Create new cell above/below | Command |  `A/B` |
| Delete cell | Command | `DD` |
| Convert cell to Markdown | Command | `M` |
| Convert cell to code | Command | `Y` |
| Show all shortcuts | Command |  `H` |




<!-- BEGIN QUESTION -->

## Question 3: Insert Cell

**Add a Markdown cell below this one, and in it write your middle name, or write "I don't have a middle name".**

```python
# The cell you create should be above this one!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 4: Delete Cell

Delete the cell below.

```python
# This should be deleted!
```

<!-- END QUESTION -->


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 2: Arithmetic Expressions and Operators

We have covered several arithmetic operators so far:

```python
(7 - 5) * 2**10 - 3 * 8 + 6 - 8 
```

Combining several `int`s without using "regular" division always results in another `int`. Below, we use the `type()` function to find the type of the evaluated expression from the previous cell.

```python
type((7 - 5) * 2**10 - 3 * 8 + 6 - 8)
```

As soon as division (`/`) or other `float`s are involved, the result is a `float`:

```python
15 / 3
```

```python
type(15 / 3)
```

We can get `int` results from division by using the **integer division** operator `//`. When the dividend (the number we're dividing) isn't evenly divisible by the divisor (the number we're dividing by), the *result is rounded down*.

```python
15 // 3
```

```python
16 // 3
```

```python
type(15 // 3)
```

We can **cast** `int`s to `float`s and vice versa by applying the `int()` or `float()` functiosn to expressions.

```python
int(15 / 3)
```

```python
float(5)
```

If our `float` is not a round integer, we lose information:

```python
int(4.1)
```

<!-- BEGIN QUESTION -->

---

## Question 5

In the cell below, **write an arithmetic expression that evaluates to your birth year *as an `int`* that uses at least four of these operators**:
- `+`
- `-`
- `*`
- `/`
- `**`

Use can also use the `int()` function.

Store the answer in "birth_year"

```python
# replace this with your expression
birth_year = ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---

## Question 6:  Pythagorean Theorem

An important data science skill is **feature engineering**, where we augment our data with an additional variable. While we leave the formal definition of "feature" to future lectures, let's use the following example to explore the benefit of using existing data to generate more information.

<br/>

Suppose the campus IT department is exploring different computer monitors for their computer lab. Computer monitor sizes are defined as diagonal measurements from corner to corner, but the IT department is concerned about the overall *monitor width* (since the monitor has to fit on their desk). They are considering monitors from the following dataset:

| Aspect Ratio | Monitor Size (inches) | Monitor Height (inches) |
| :--- | :---- | :---- |
| 16:9 | 20 | 9.8 | 
| 16:9 | 22 | 10.8 |
| 16:9 | 24 | 11.8 |

The data above include aspect ratio (all widescreen, 16:9), monitor size in inches, and monitor height in inches–but no monitor width. As a data scientist, you decide to generate another feature for each listed monitor in the data by computing the associated monitor width.

In the cell below, **write an arithmetic expression that, when evaluated, computes the monitor width for the 22" monitor size** listed in the table above.

**Hints**:
* Recall the Pythagorean theorem for a right triangle: $a^2 + b^2 = c^2$, where $a$, $b$, and $c$ are the lengths of the first leg, second leg, and hypotenuse, respectively.
* We didn't explicitly cover how to take the square root of a number, but recall that in mathematics, the square root function $\sqrt{x}$ is equivalent to exponentiating with the exponent $\frac{1}{2}$ – that is, $\sqrt{x} = x^{\frac{1}{2}}$. That means we can use the exponentiation operator `**` to compute square roots, too. For example, `121 ** 0.5` evaluates to `11.0`.
* Like in math, may need to use parentheses `(` and `)` for order of operations.
* There are two valid ways of computing the monitor width. Each method uses values from two out of the three available columns.
* Feel free to validate your result with a screen size calculator, e.g., [http://screen-size.info/](http://screen-size.info/).

```python
# replace this with your expression
```

<!-- END QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 3: Names

In this part, explore a Tuberculosis dataset. Here's the data:

Tuberculosis data. CDC MMWR [source](https://www.cdc.gov/mmwr/volumes/71/wr/mm7112a1.htm?s_cid=mm7112a1_w#T1_down)

| U.S. jurisdiction | 2019<br/># cases | 2020<br/># cases | 2021 <br/> # cases | 2019<br/>incidence | 2020<br/>incidence | 2021<br/>incidence |
| --- | --- | --- | --- | --- | --- | --- |
| **Total** | **8,900** | **7,173** | **7,860** | **2.71** | **2.16** | **2.37** | 
| Alabama | 87 | 72 | 92 | 1.77 | 1.43 | 1.83 |
| ... | ... | ... | ... | ... | ... | ... |
| California | 2,111 | 1,706 | 1,750 | 5.35 | 4.32 | 4.46 |
| ... | ... | ... | ... | ... | ... | ... |


U.S. Census population estimates ([2019 source](https://www.census.gov/data/tables/time-series/demo/popest/2010s-state-total.html), [2021-2021 source](https://www.census.gov/data/tables/time-series/demo/popest/2020s-state-total.html))

|  | 2019 population | 2020 population | 2021 population | 
| --- | --- | --- | --- |
| **Total** | **328,239,523** | **331,501,080** | **331,893,745** |
| Alabama | 4,903,185 | 5,024,803 | 5,039,877 |
| California | 39,512,223 | 39,499,738 | 39,237,836 |


Recall that the Tuberculosis (TB) incidence is defined as the number of TB cases per 100,000 persons in a particular population:

$$ \text{TB incidence} = \frac{\text{\# TB cases}}{\text{population}} \times 100000$$

We computed the tuberculosis (TB) incidence of 2020 Using Python's arithmetic expressions, this value can be written as:


```python
7173/(331501080/100000)
```

We'll go one step further in the cell below by binding (i.e., assigning) a name to the value returned from evaluating this arithmetic expression. Below, note the name `incidence_2020` clearly indicates what year's incidence we are computing.

```python
# just run this cell
incidence_2020 = 7173/(331501080/100000)
incidence_2020
```

**Notice!** You'll often see the above *syntax* (i.e., program code structure) in Jupyter notebooks. It is effectively evaluating **two** Python statements:

1. ```incidence_2020 = 7173/(331501080/100000)```

    Assign the name `incidence_2020` to the result of evaluating the arithmetic expression `7173/(331501080/100000)`.<br/><br/>

1. ```incidence_2020```

    Evaluate the expression `incidence_2020`. Since `incidence_2020` is a name, the result of this arithmetic expression is the value bound to the name `incidence_2020`.

**Remember!** Jupyter code cells output only the value of the *last* evaluated expression. So this two-step syntax is useful to bind a name to a value and output the value itself.

---

## Question 7: Changes in Percentage

Consider this line in the abstract of the CDC MMWR [report](https://www.cdc.gov/mmwr/volumes/71/wr/mm7112a1.htm):
> Reported TB incidence (cases per 100,000 person) increased 9.4%, from 2.2 during 2020 to 2.4 during 2021...


The 2.2 TB incidence in the quote above is the same incidence we computed above, *rounded* to the first decimal place (i.e., to the nearest tenth). For the purposes of this exercise, we will not work with any rounded numbers, and then qualitatively compare our computations to the quote above.

### Question 7

**Assign the name `incidence_2021` to the 2021 tuberculosis incidence. Note we use the same structure as explained above to bind the name and output its value in a single code cell.**

*Hint*: Use the tables above.

**Note:** On labs, questions that ask you to write code will usually be followed with a cell that says something like `otter.grade('q7')` that you'll need to run in order for your work to be graded. 

```python
incidence_2021 = ...
incidence_2021
```

```python
grader.check("q7")
```

### Question 8
Compute percent change in national TB cases from 2020 to 2021.

*Hint*: Use the tables above.

```python
# Calculate percent change in total TB cases from 2020 to 2021

percent_change_cases = ...
percent_change_cases
```

```python
grader.check("q8")
```

### Question 9
Compute relative change ratios year-over-year.

*Hint*: Use the tables above.

```python
# Given a list of values (e.g., case counts or incidences), return a list of 
# relative year-over-year changes rounded to 3 decimal places.
# For input [a, b, c], return [b/a, c/b].

def relative_changes(values):
    ...

relative_changes([8900, 7173, 7860])
```

```python
grader.check("q9")
```

<hr>

## Congratulations! 😇

Try submitting this assignment to Gradescope. 

1. Save your notebook 
2. Run the last cell 
3. Get the zip file and upload it to Gradescope 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q7 (points: 4)
```
>>> type(incidence_2021) == float
True
```
```
>>> np.isclose(incidence_2021, 2.4, 0.1)
np.True_
```


Tests for q8 (points: 4)
```
>>> type(percent_change_cases) == float
True
```
```
>>> type(percent_change_cases) == float
True
```
```
>>> np.isclose(percent_change_cases, 9.5, 0.1)
np.True_
```


Tests for q9 (points: 4)
```
>>> np.allclose(relative_changes([8900, 7173, 7860]), [0.806, 1.096], atol=0.001)
True
```
```
>>> np.allclose(relative_changes([100, 200, 400]), [2.0, 2.0], atol=0.001)
True
```
```
>>> np.allclose(relative_changes([100, 150, 75, 90]), [1.5, 0.5, 1.2], atol=0.001)
True
```
